```python
# Initialize Otter
import otter
grader = otter.Notebook("wordle_lab.ipynb")
```

# Wordle!

#### Authors:
v1.0 (Fall 2022): Andy Dong

v1.1 (Spring 2025): Lance Mathias

v1.2 (Fall 2026): Rex Ouyang

## Setup
This lab runs on JupyterHub or on your own machine. The first time you run it, the notebook builds a lookup table of Wordle color patterns (about 30 MB, usually under a minute) and caches it in the `checkpoint/` folder, so later runs start right away.

**How this lab is organized**
- Sections marked **(Read Only)** contain provided code and explanations. Read them and run their cells; the provided code cells are locked and you don't need to change them.
- **Tasks** are functions for you to implement wherever you see `...`. Each task is followed by a `grader.check(...)` cell that runs its tests.
- **Written Questions** are answered in the accompanying Gradescope assignment, along with any plots they ask for.
- The **Wordle Bot** section at the end is optional and not graded.

You don't need to know any information theory for this lab. Everything is built from conditional probability and expectation.

## Introduction

In the game of [Wordle](https://www.nytimes.com/games/wordle/index.html), your goal is to deduce a 5-letter secret word that the game has generated using as few guesses as possible. After each guess, the game will tell you for each letter in your guess, whether the letter is

-  in the secret word and in the right spot (indicated by green),
-  in the secret word but in the wrong spot (indicated by yellow),
-  or not in the secret word (indicated by gray).

Take a look at the example game below where the secret word is "scald".

<div>
<img src="examples/ex-0.png" width="500"/>
</div>

Suppose we want to write a program that solves the game in as few guesses as possible. In this lab, you will build a greedy solver using conditional probability: after every guess, work out which words the secret could still be, then pick the next guess that is expected to narrow them down the most. It's recommended that you play a game of Wordle first to get familiar with the rules.

```python
# If modules have not been installed yet:

%pip install -r requirements.txt

# Restart kernel afterwards, if needed
```

```python
import math
import os
import random
from functools import lru_cache

import matplotlib.pyplot as plt
import numpy as np
from scipy import stats
from tqdm import tqdm

rng_seed = 0
random.seed(rng_seed)  # makes the random samples in this notebook reproducible
```

## Modeling the Game

Here's some setup and simplifying assumptions about the game. First, we have obtained the list of 2315 words that can be chosen as the secret word. The secret word $X$ is modeled to be a random variable that has uniform distribution over these possible words. Second, your guess cannot be any arbitrary combination of 5 letters but needs to be from a list of 12972 valid guesses (which, of course, contains the 2315 possible secret words). Third, any revealed hints don't need to be used in subsequent guesses. With these, let's dive in!

Run the following cell to read in the list of possible secret words and the list of allowed guesses.

```python
with open('data/possible_words.txt') as file:
    possible_words = [line.rstrip() for line in file]
with open('data/allowed_guesses.txt') as file:
    allowed_guesses = [line.rstrip() for line in file]

print('Number of possible_words:', len(possible_words))
print('Number of allowed_guesses:', len(allowed_guesses))
```

Let's print out some of these words to take a look. `possible_words` consists of words used more frequently in the English language while `allowed_guesses` contains some words that we don't see in day-to-day life.

```python
print('Some possible secret words:')
print(*random.sample(possible_words, 20), '\n')
print('Some allowed guesses:')
print(*random.sample(allowed_guesses, 20))
```

## Implementing the Game Logic (Read Only)

We then implement the logic of the game by first computing the pattern that we will see for a given pair of guess and secret word. The function `compute_pattern` takes in a guess and a secret word, and returns an integer tuple of length 5 that represents the color pattern, where an entry of 2 represents green in the corresponding position, 1 represents yellow, and 0 represents gray.

There's an edge case that we will need to consider, which is when the guess and/or the secret word contains multiple of the same letter. The rule here is that green (match in the correct spot) always has the highest priority, yellow (match in the wrong spot) prioritizes letters in earlier positions of your guess, and each letter in the secret word can only correspond to one letter in the guess. Take a look at the examples below.

- If the secret word is "abide" and the guess is "three", the pattern would be (0, 0, 0, 0, 2). Position 5 is an exact match so it has the highest priority to be green. Position 4 is indeed a match in the wrong spot, but since the "e" in "abide" is already taken, it cannot match to that same "e" again.

![title](examples/ex-1.png)

- If the secret word is "abide" and the guess is "drama", the pattern would be (1, 0, 1, 0, 0). Position 3 is a yellow match that has priority over the potential yellow match at position 5.

![title](examples/ex-2.png)

```python
@lru_cache(maxsize=None)
def compute_pattern(guess, answer):
    # Returns a length 5 tuple

    pattern = [0, 0, 0, 0, 0]
    taken = [False, False, False, False, False]

    # Green pass
    for i in range(5):
        if guess[i] == answer[i]:
            # If it's an exact match, color it green, and mark it as taken
            # so that the yellow pass doesn't match to it again
            pattern[i] = 2
            taken[i] = True

    # Yellow pass
    for i in range(5):
        if pattern[i] == 2:
            # If a spot is already colored green, we skip it
            continue
        query = guess[i]
        for j in range(5):
            if query == answer[j] and taken[j] is False:
                # If there is a misplaced match that is not taken by the
                # green pass or a previous yellow pass, we color it yellow
                # and mark it as taken
                pattern[i] = 1
                taken[j] = True
                break

    return tuple(pattern)
```

## Pattern Lookup Table (Read Only)

The solver needs the color pattern for millions of (guess, secret word) pairs, so we compute all of them once and store them in a table. Each pattern is encoded as an integer in $[0, 243)$ by reading it as a base-3 number, e.g. (2, 1, 0, 0, 2) $\to$ 191, and the table is a matrix with one row per allowed guess and one column per possible secret word (12972 $\times$ 2315, about 30 MB).

**In the rest of the lab, use `pattern_table[guess][answer]` instead of `compute_pattern(guess, answer)`.** It returns the same tuple but is much faster. It also works for words outside the precomputed lists by falling back to `compute_pattern`. You don't need to understand how the table is built.

```python
PATTERN_POWERS = 3 ** np.arange(4, -1, -1)  # [81, 27, 9, 3, 1]

def pattern_to_code(pattern):
    # length-5 tuple -> integer in [0, 243), e.g. (2, 1, 0, 0, 2) -> 191
    return int(np.dot(pattern, PATTERN_POWERS))

@lru_cache(maxsize=None)
def code_to_pattern(code):
    # integer in [0, 243) -> length-5 tuple
    digits = []
    for _ in range(5):
        digits.append(code % 3)
        code //= 3
    return tuple(reversed(digits))

def build_pattern_matrix(guesses, answers, chunk_size=512):
    # Vectorized equivalent of compute_pattern over all (guess, answer) pairs.
    # Returns a (len(guesses), len(answers)) uint8 matrix of base-3 pattern codes.
    g = np.array([[ord(c) - ord('a') for c in w] for w in guesses], dtype=np.uint8)
    a = np.array([[ord(c) - ord('a') for c in w] for w in answers], dtype=np.uint8)
    matrix = np.zeros((len(guesses), len(answers)), dtype=np.uint8)

    for start in tqdm(range(0, len(guesses), chunk_size)):
        gc = g[start:start + chunk_size]
        # eq[i, j, p, q]: does letter p of guess i equal letter q of answer j?
        eq = gc[:, None, :, None] == a[None, :, None, :]
        colors = np.zeros(eq.shape[:3], dtype=np.uint8)

        # Green pass: exact matches take the answer letter out of play
        for p in range(5):
            green = eq[:, :, p, p].copy()
            colors[:, :, p][green] = 2
            eq[:, :, :, p] &= ~green[:, :, None]

        # Yellow pass: earlier guess positions have priority, and each answer
        # letter can only be matched once
        for p in range(5):
            for q in range(5):
                match = eq[:, :, p, q] & (colors[:, :, p] == 0)
                colors[:, :, p][match] = 1
                eq[:, :, :, q] &= ~match[:, :, None]

        matrix[start:start + chunk_size] = colors @ PATTERN_POWERS

    return matrix

pattern_matrix_path = os.path.join('checkpoint', 'pattern_matrix.npy')
expected_shape = (len(allowed_guesses), len(possible_words))

pattern_matrix = None
if os.path.exists(pattern_matrix_path):
    pattern_matrix = np.load(pattern_matrix_path)
if pattern_matrix is None or pattern_matrix.shape != expected_shape:
    # build (or rebuild, if the word lists changed) and cache the table
    pattern_matrix = build_pattern_matrix(allowed_guesses, possible_words)
    os.makedirs('checkpoint', exist_ok=True)
    np.save(pattern_matrix_path, pattern_matrix)

guess_index = {w: i for i, w in enumerate(allowed_guesses)}
answer_index = {w: i for i, w in enumerate(possible_words)}

def get_pattern(guess, answer):
    # Fast table lookup; falls back to computing directly for words outside
    # the precomputed lists
    gi, ai = guess_index.get(guess), answer_index.get(answer)
    if gi is None or ai is None:
        return compute_pattern(guess, answer)
    return code_to_pattern(int(pattern_matrix[gi, ai]))

class _PatternTable:
    # pattern_table[guess][answer] lookups on top of the numpy matrix
    class _Row:
        def __init__(self, guess):
            self.guess = guess
        def __getitem__(self, answer):
            return get_pattern(self.guess, answer)
    def __getitem__(self, guess):
        return self._Row(guess)

pattern_table = _PatternTable()

# Sanity check: the table must agree with compute_pattern
_rng = random.Random(rng_seed)
for _ in range(2000):
    _g, _a = _rng.choice(allowed_guesses), _rng.choice(possible_words)
    assert pattern_table[_g][_a] == compute_pattern(_g, _a)
print('Pattern table ready:', pattern_matrix.shape, '- matches compute_pattern on 2000 random pairs')
```

## Guess Strategy

We use the term *alphabet* of a random variable to mean the set of values it takes with positive probability, i.e. its support (not to be confused with the English alphabet). For example, the alphabet of $X$ is the set of 2315 possible secret words.

Let $Y_i$ be the color pattern we observe after our $i$-th guess. Fix a time index $t$ and suppose we have already observed $Y_1 = y_1, \ldots, Y_{t-1} = y_{t-1}$. Everything we know about the secret word at that point is captured by the conditional distribution of $X$ given these observations. Let $X_t$ be a random variable with that distribution, and let $\mathcal{A}_t$ be its alphabet. At $t = 1$ nothing has been observed yet, so $X_1 = X$ and $\mathcal{A}_1$ is all 2315 possible words.

If we guess the word $k$ at time $t$, the pattern we will see is $Y_{t,k} = f_k(X_t)$, where $f_k(x)$ = `compute_pattern(k, x)`. Note that the index $k$ is a word we guess, and that $Y_{t,k}$ is a *deterministic* function of the secret word: all of its randomness comes from not knowing $X_t$.

<!-- BEGIN QUESTION -->

**Written Question 1:** Suppose $X_t$ is uniform over its alphabet $\mathcal{A}_t$. We guess $k$ and observe $Y_{t,k} = y$.

1. Use Bayes' rule to show that, conditioned on $Y_{t,k} = y$, the secret word is uniform over $\mathcal{A}_{t+1} = \{x \in \mathcal{A}_t : f_k(x) = y\}$.
2. Conclude that $X_t$ is uniform over its alphabet at every time step $t$.
3. Suppose instead that the secret word were *not* uniform to begin with (e.g. common words were more likely). Would $X_t$ still be uniform over its alphabet? What can you say about its distribution instead?

Feel free to use this space to draft your answer. When you finish, remember to paste your response and any relevant plots/figures **into the accompanying Gradescope assignment**.

<!-- END QUESTION -->

In other words, each guess simply *narrows down* the alphabet: we keep exactly the words that are consistent with the pattern we saw, and they all remain equally likely.

**Task 1:** Implement `filter_alphabet`, which takes in the current alphabet (a list of words), the word we guessed, and the color pattern we observed, and returns a list of the words in the alphabet that are consistent with the observation. This is the alphabet $\mathcal{A}_{t+1}$ from Written Question 1. For example, `filter_alphabet(['shape', 'shake', 'shame'], 'shake', (2, 2, 2, 0, 2))` should return `['shape', 'shame']` (in any order).

```python
def filter_alphabet(alphabet, guess, color_pattern):
    # Returns a list of the words in `alphabet` that would have produced
    # `color_pattern` if they were the secret word
    color_pattern = tuple(color_pattern)
    ...
```

```python
grader.check("filter_alphabet")
```

## Choosing a Guess

Before we make guess $k$, the pattern $Y_{t,k}$ is random. Guessing $k$ splits the alphabet $\mathcal{A}_t$ into groups of words that share the same pattern, and after the guess we are left with exactly one of those groups. A good guess splits the alphabet into many small groups; a bad guess leaves most of the words together in one big group.

**Task 2:** Implement `divide_alphabet`, which takes in a guess and the current alphabet and splits the alphabet into these groups. Namely, it returns a dictionary that maps each possible color pattern to the list of secret words such that `pattern_table[guess][secret_word]` is that color pattern. For example, if `guess` is "shake" and `alphabet` is {shape, shake, shame}, then the function should return the mapping (2,2,2,2,2)$\to${shake}, (2,2,2,0,2)$\to${shape, shame}.

```python
def divide_alphabet(guess, alphabet):
    pattern_to_subgroup = {}
    ...
    return pattern_to_subgroup
```

```python
grader.check("divide_alphabet")
```

**Task 3:** Since $X_t$ is uniform over $\mathcal{A}_t$, the probability of observing each pattern is proportional to the size of its group: $P(Y_{t,k} = y) = n_y / N$, where $N = |\mathcal{A}_t|$ and $n_y$ is the number of words in the group for pattern $y$.

Implement `prob_dist`, which takes in the output of `divide_alphabet` and returns the probability distribution over the set of possible patterns. For the same example above, this function would take in the mapping (2,2,2,2,2)$\to${shake}, (2,2,2,0,2)$\to${shape, shame} and return $\left[\frac{1}{3}, \frac{2}{3}\right]$. This is the PMF of $Y_{t,k}$.

```python
def prob_dist(pattern_groups):
    # returns a probability distribution in the form of a list
    # for example, if the probability distribution is
    # P(pattern_1) = 0.2, P(pattern_2) = 0.3, P(pattern_3) = 0.5,
    # this function should return [0.2, 0.3, 0.5]. Order doesn't matter
    ...
```

```python
grader.check("prob_dist")
```

### Scoring a Guess

To compare guesses, we summarize each pattern distribution with a single number. We'll look at two such scores.

**Score 1: expected number of remaining words.** After guessing $k$, the new alphabet is the group of whichever pattern we actually observe, so its size $|\mathcal{A}_{t+1}| = n_{Y_{t,k}}$ is a random variable. A natural goal is to choose the guess that makes it small on average.

<!-- BEGIN QUESTION -->

**Written Question 2:** Let $p_y = P(Y_{t,k} = y) = n_y / N$.

1. Show that $E\big[|\mathcal{A}_{t+1}|\big] = N \sum_y p_y^2$.
2. Let $X$ and $X'$ be two independent words, each drawn uniformly from $\mathcal{A}_t$. Show that $\sum_y p_y^2 = P\big(f_k(X) = f_k(X')\big)$, i.e. the probability that guess $k$ fails to tell $X$ and $X'$ apart. Use this to explain in one sentence what a guess with small $E\big[|\mathcal{A}_{t+1}|\big]$ is good at.
3. What is $E\big[|\mathcal{A}_{t+1}|\big]$ if every word in $\mathcal{A}_t$ gives a different pattern? If every word gives the same pattern?

Feel free to use this space to draft your answer. When you finish, remember to paste your response and any relevant plots/figures **into the accompanying Gradescope assignment**.

<!-- END QUESTION -->

**Task 4:** Implement `expected_remaining`, which takes in a pattern distribution (the output of `prob_dist`) and the current alphabet size $N$, and returns $E\big[|\mathcal{A}_{t+1}|\big]$.

```python
def expected_remaining(dist, num_words):
    # Returns E[|A_{t+1}|], the expected number of words left after the guess,
    # given its pattern distribution `dist` and the current alphabet size N = num_words
    ...
```

```python
grader.check("expected_remaining")
```

**Score 2: expected number of halvings (Read Only).** Score 1 counts every eliminated word equally, so cutting the alphabet from 1000 words down to 500 counts 125 times as much as cutting it from 8 down to 4. Yet each of those cuts halves the alphabet, and each gets us about equally closer to the answer. An alternative is to count halvings instead: going from $N$ words down to $n_y$ words halves the alphabet $\log_2(N / n_y) = \log_2(1/p_y)$ times. So the expected number of halvings from guess $k$ is

$$H(Y_{t,k}) = E\left[\log_2 \frac{N}{|\mathcal{A}_{t+1}|}\right] = \sum_y p_y \log_2 \frac{1}{p_y}.$$

This quantity is called the **entropy** of $Y_{t,k}$, and it is measured in **bits** (one bit = one halving). You won't need any other properties of entropy in this lab. For example, pinning down one of 2315 equally likely secret words from scratch takes $\log_2 2315 \approx 11.18$ halvings. The function below computes the entropy of a distribution in the same list format as the output of `prob_dist`.

```python
def entropy(dist):
    # Expected number of halvings (in bits) for a pattern distribution with no zero entries
    return sum(-p * math.log2(p) for p in dist)
```

## The Solver (Read Only)

Finding the best guess means scoring every one of the 12972 allowed guesses against the current alphabet. Doing that with `divide_alphabet` would take minutes, so the code below computes the same pattern counts for all guesses at once using the lookup table.

- `guess_entropies(alphabet, guesses)` and `guess_expected_remaining(alphabet, guesses)` compute Score 2 and Score 1 for every word in `guesses`. They should agree exactly with `entropy(prob_dist(divide_alphabet(guess, alphabet)))` and with your `expected_remaining`; the check after this cell tests this.
- `find_best_n_guesses(alphabet, guesses, n, score)` returns the `n` best guesses along with their scores, where `score` is `'entropy'` (most halvings first, the default) or `'remaining'` (fewest expected remaining words first).
- `find_best_guess(alphabet, guesses, score)` returns just the best guess.

When several guesses tie, the solver prefers one that is in the alphabet, since it might be the secret word itself.

```python
def pattern_codes(guesses, answers):
    # (len(guesses), len(answers)) matrix of pattern codes: read from the
    # lookup table where possible, computed directly for any other words
    rows = np.array([guess_index.get(g, -1) for g in guesses], dtype=np.int64)
    cols = np.array([answer_index.get(a, -1) for a in answers], dtype=np.int64)
    known_rows, known_cols = rows >= 0, cols >= 0
    if known_rows.all() and known_cols.all():
        return pattern_matrix[np.ix_(rows, cols)]
    codes = np.empty((len(guesses), len(answers)), dtype=np.uint8)
    codes[np.ix_(known_rows, known_cols)] = pattern_matrix[np.ix_(rows[known_rows], cols[known_cols])]
    for i in np.nonzero(~known_rows)[0]:
        codes[i, :] = [pattern_to_code(compute_pattern(guesses[i], a)) for a in answers]
    for j in np.nonzero(~known_cols)[0]:
        codes[:, j] = [pattern_to_code(compute_pattern(g, answers[j])) for g in guesses]
    return codes

def pattern_counts(alphabet, guesses):
    # counts[i, c] = number of words in alphabet that give pattern code c for guesses[i]
    codes = pattern_codes(list(guesses), list(alphabet)).astype(np.int64)
    codes += np.arange(len(guesses))[:, None] * 243  # offset so each guess gets its own bins
    counts = np.bincount(codes.ravel(), minlength=len(guesses) * 243)
    return counts.reshape(len(guesses), 243)

def guess_entropies(alphabet, guesses):
    # Score 2, H(Y_{t,k}), for every guess k
    probs = pattern_counts(alphabet, guesses) / len(alphabet)
    plogp = np.zeros_like(probs)
    nonzero = probs > 0
    plogp[nonzero] = probs[nonzero] * np.log2(probs[nonzero])
    return -plogp.sum(axis=1)

def guess_expected_remaining(alphabet, guesses):
    # Score 1, E[|A_{t+1}|], for every guess k
    counts = pattern_counts(alphabet, guesses).astype(np.float64)
    return (counts ** 2).sum(axis=1) / len(alphabet)

def find_best_n_guesses(alphabet, guesses, n=5, score='entropy'):
    # The n best guesses as (word, score) pairs, best first
    guesses = list(guesses)
    if score == 'entropy':
        values = guess_entropies(alphabet, guesses)
        goodness = values
    elif score == 'remaining':
        values = guess_expected_remaining(alphabet, guesses)
        goodness = -values
    else:
        raise ValueError("score must be 'entropy' or 'remaining'")
    alpha_set = set(alphabet)
    in_alpha = np.array([g in alpha_set for g in guesses], dtype=np.int8)
    # rounding makes floating-point near-ties exact, so ties go to words in the alphabet
    order = np.lexsort((-in_alpha, -np.round(goodness, 9)))[:n]
    return [(guesses[i], float(values[i])) for i in order]

def find_best_guess(alphabet, guesses, score='entropy'):
    if len(alphabet) == 0 or len(guesses) == 0:
        return None
    return find_best_n_guesses(alphabet, guesses, n=1, score=score)[0][0]
```

```python
grader.check("solver_agreement")
```

## Best Wordle Opener (Read Only)

Here is the question that Wordle lovers are looking for: What is the best opening guess under our scheme? (**Important Note:** this is the best opener given our assumptions that the secret word is picked *uniformly at random* over the 2315 possible secret words, and that we are using a greedy algorithm that maximizes the expected number of halvings of each guess. This is not guaranteed to be the truly optimal guess.)

```python
best_opener = find_best_guess(possible_words, allowed_guesses)
print('The best Wordle opener is:', best_opener)
```

The best opener is "soare", which is an obsolete form of sore ("A young hawk") according to the Collins dictionary. Let's plot the distribution of patterns we observe if we open with "soare", and compare it with opening with "speed" (a mediocre opener) and with "qajaq" (an alternative spelling of kayak, and the worst opener). Since there's no natural ordering of patterns, we order the patterns on the x-axis by decreasing probability.

These plots use your `divide_alphabet`, `prob_dist`, and `expected_remaining`.

```python
def plot_pattern_distribution(guess):
    # Bar chart of the opener's pattern distribution, patterns sorted by decreasing probability
    dist = prob_dist(divide_alphabet(guess, possible_words))
    print(f'{guess.upper()}: {len(dist)} possible patterns, '
          f'{entropy(dist):.2f} expected halvings (bits) out of {math.log2(len(possible_words)):.2f}, '
          f'{expected_remaining(dist, len(possible_words)):.1f} expected remaining words')
    dist.sort(reverse=True)
    plt.bar(range(1, len(dist) + 1), dist)
    plt.title(f'Distribution of Patterns for {guess.upper()} Opening')
    plt.xlabel('Pattern Number')
    plt.ylabel('Probability')
    plt.show()
```

```python
plot_pattern_distribution('soare')
```

```python
plot_pattern_distribution('speed')
```

```python
plot_pattern_distribution('qajaq')
```

The "qajaq" opener is so bad that only a few patterns have a positive probability of appearing. Note that the horizontal axis of the last graph is on a different scale from the previous graphs.

The two scores don't always agree on which guess is best. The cell below lists the top 10 openers under each score.

```python
for score, unit in [('entropy', 'expected halvings (bits)'), ('remaining', 'expected remaining words')]:
    print(f'Top 10 openers by {score}:')
    for word, value in find_best_n_guesses(possible_words, allowed_guesses, n=10, score=score):
        print(f'  {word}: {value:.3f} {unit}')
    print()
```

<!-- BEGIN QUESTION -->

**Written Question 3:** 1. Using the three plots, explain why a flatter pattern distribution is better under both scores. Why is "qajaq" such a bad opener?
2. Which opener minimizes the expected number of remaining words? Which of the two scores would you rather use to pick guesses, and why? There is no single right answer; argue from what each score rewards.

Feel free to use this space to draft your answer. When you finish, remember to paste your response and any relevant plots/figures **into the accompanying Gradescope assignment**.

<!-- END QUESTION -->

## Some Useful Optimizations

Two tweaks make the solver faster and better. First, the opening guess is always made from the same alphabet (all of `possible_words`), so it is always `best_opener`, and there is no need to recompute it every game.

Second, when only a few words remain, it is better to guess one of them than a more informative word that can't be the answer. The next question makes this precise.

<!-- BEGIN QUESTION -->

**Written Question 4:** Let $G$ be the number of guesses we still need to make until we win, including the winning guess. Assume the secret word is uniform over the current alphabet.

1. The alphabet is {"price", "pride"}. Compute $E[G]$ if we first guess "camel" (which tells the two apart), and $E[G]$ if we first guess "price".
2. Now suppose the alphabet has 3 words. Compute $E[G]$ if we first guess (a) a word in the alphabet whose pattern tells the other two apart, (b) a word in the alphabet whose pattern cannot tell the other two apart, and (c) a word outside the alphabet that tells all three apart. In each case, assume we play optimally afterwards.
3. Based on your answers, explain why it makes sense to restrict our guesses to the alphabet when $|\mathcal{A}_t| \leq 3$.

Feel free to use this space to draft your answer. When you finish, remember to paste your response and any relevant plots/figures **into the accompanying Gradescope assignment**.

<!-- END QUESTION -->

**Task 5:** Implement `find_best_guess_optimized`, which takes in the current alphabet and returns the next word to guess using the default entropy score:
- If this is the opening guess (the alphabet is all of `possible_words`), return `best_opener` directly.
- If only one word remains, return it.
- If 2 or 3 words remain, return the best guess from within the alphabet, i.e. `find_best_guess(alphabet, alphabet)`.
- Otherwise, return the best guess from all of `allowed_guesses`.

```python
def find_best_guess_optimized(alphabet):
    ...
```

```python
grader.check("find_best_guess_optimized")
```

## Performance Test

Let's now sample some secret words and watch our solver find the secret word in action! The provided function `create_wordle_game` takes in a secret word and returns a function that plays the role of the game: given a guess, it returns the color pattern.

```python
def create_wordle_game(true_answer):
    def wordle_game(guess):
        # takes in a guess and outputs the pattern
        return compute_pattern(guess, true_answer)

    return wordle_game
```

**Task 6:** Implement `play_wordle`, which takes in a game created by `create_wordle_game` and plays it until it wins: guess with `find_best_guess_optimized`, then narrow down the alphabet with `filter_alphabet` using the pattern the game returns. It should return the number of guesses used, including the final correct guess. If `print_guesses` is `True`, also print each guess and its pattern.

```python
def play_wordle(wordle_game, print_guesses=False):
    alphabet = possible_words
    num_guesses = 0
    ...
    return num_guesses
```

```python
for _ in range(10):
    true_answer = random.choice(possible_words)
    print('Secret word:', true_answer)
    wordle_game = create_wordle_game(true_answer)
    play_wordle(wordle_game, print_guesses=True)
    print()
```

Let's take a look at the distribution of the number of guesses needed by sampling 300 secret words.

```python
num_guesses = []
num_samples = 300
for true_answer in tqdm(random.sample(possible_words, num_samples)):
    num_guesses.append(play_wordle(create_wordle_game(true_answer)))
```

```python
bins = list(range(1, 8))
plt.hist(num_guesses, bins, alpha=0.5, edgecolor='black', align='left')
plt.xlabel('Number of Guesses')
plt.xticks(list(range(1, 8)))
plt.title('Wordle Solver Performance')
plt.show()
```

If your implementation is correct, your Wordle solver takes fewer guesses on average than American players do.

```python
print(f'Average number of Wordle guesses using entropy: {round(sum(num_guesses)/len(num_guesses), 3)}')
print('Average number of Wordle guesses for Americans: 3.92')
```

```python
grader.check("play_wordle")
```

## How Good Is the Solver?

Our solver is deterministic, so its average number of guesses over all 2315 possible secret words is a fixed number $\mu$. Above, we estimated $\mu$ with the sample mean $\bar{G}$ of 300 randomly chosen secret words. How confident can we be that $\mu < 3.92$?

**Task 7:** Implement `confidence_interval`, which takes in a list of samples and returns a tuple `(low, high)`: the approximate confidence interval for the mean given by the Central Limit Theorem,
$$\bar{G} \pm z \frac{s}{\sqrt{n}},$$
where $n$ is the number of samples, $s$ is the sample standard deviation (with $n - 1$ in the denominator, i.e. `np.std(samples, ddof=1)`), and $z$ is chosen so that $P(-z \leq Z \leq z)$ equals `confidence` for $Z \sim \mathcal{N}(0, 1)$. You may find `stats.norm.ppf` helpful.

```python
def confidence_interval(samples, confidence=0.95):
    # Returns (low, high), the CLT confidence interval for the mean of `samples`
    ...
```

```python
grader.check("confidence_interval")
```

```python
low, high = confidence_interval(num_guesses, confidence=0.95)
print(f'95% confidence interval for the average number of guesses: ({low:.3f}, {high:.3f})')
```

<!-- BEGIN QUESTION -->

**Written Question 5:** 1. Report your 95% confidence interval. Can you conclude that the solver beats the 3.92 average of American players?
2. Chebyshev's inequality gives a confidence interval without relying on the CLT. Assuming the standard deviation $\sigma$ of the number of guesses is known, how wide is the 95% interval from Chebyshev's inequality compared to the CLT interval? Plugging in $s$ for $\sigma$, would your conclusion in part 1 change?
3. The 300 secret words were sampled *without* replacement, so the samples are not independent. Why is the CLT interval still reasonable here? What would happen if we sampled all 2315 words?

Feel free to use this space to draft your answer. When you finish, remember to paste your response and any relevant plots/figures **into the accompanying Gradescope assignment**.

<!-- END QUESTION -->

## Wordle Bot (Optional)

Finally, let's turn the solver into an assistant for real games! This section is optional and not graded.

Implement the class `WordleBot` below. It keeps track of the words that the secret could still be, suggests guesses, and updates when you tell it which pattern the game showed:
- `suggest()` prints how many words remain (listing them if there are only a few), prints the top `num_suggestions` guesses with their scores, and returns the best one.
- `observe(word, pattern)` narrows down the alphabet using the word you actually guessed and the pattern the game showed, then calls and returns `suggest()`.
- `restart()` resets the bot for a new game.

Some ideas for extending the bot:
- Show each suggestion's score under both `'entropy'` and `'remaining'`.
- **Uncommon answers.** The NYT has occasionally picked answers that are valid guesses but aren't in `possible_words`, like "pshaw". If the bot's only candidates are `possible_words`, it runs out of words in that case. Keep a second candidate set drawn from `allowed_guesses` and fall back to it when no word from `possible_words` fits. (`pattern_table`, `filter_alphabet`, and the solver functions work for any 5-letter words; they're just slower for words outside `possible_words`.)

```python
class WordleBot:
    def __init__(self, num_suggestions=10):
        self.num_suggestions = num_suggestions
        self.restart()

    def suggest(self):
        # TODO: print the remaining words (or how many there are) and the top
        # self.num_suggestions guesses with their scores; return the best guess
        ...

    def observe(self, word, pattern):
        # after a guess, feed the pattern to the bot to update
        # then, the bot suggests a word to guess
        assert len(word) == len(pattern) == 5
        # TODO: update the alphabet according to the observation,
        # then call and return self.suggest()
        ...

    def restart(self):
        # TODO: re-initialize the bot for a new game
        ...
```

Try your bot on an example game. The secret word here is "whine".

```python
bot = WordleBot()
bot.suggest()
```

```python
bot.observe('soare', (0, 0, 0, 0, 2))
```

```python
bot.observe('guilt', (0, 0, 2, 0, 0))
```

```python
bot.observe('chide', (0, 2, 2, 0, 2))
```

```python
bot.observe('whine', (2, 2, 2, 2, 2))
```

Now an uncommon answer: the secret word here is "pshaw", which is not in `possible_words`. Your bot will only handle this game if you implemented the fallback extension.

```python
bot.restart()
```

```python
bot.observe('soare', (1, 0, 1, 0, 0))
```

```python
bot.observe('linty', (0, 0, 0, 0, 0))
```

```python
bot.observe('chump', (0, 1, 0, 0, 1))
```

```python
bot.observe('hasps', (1, 1, 1, 1, 0))
```

```python
bot.observe('pshaw', (2, 2, 2, 2, 2))
```

Tests for confidence_interval (points: None)
```
>>> low, high = confidence_interval([1, 2, 3, 4, 5])
>>> assert math.isclose((low + high) / 2, 3), 'The interval should be centered at the sample mean'
>>> assert math.isclose(high - low, 2 * 1.959963984540054 * math.sqrt(2.5 / 5), rel_tol=1e-6), 'Wrong interval width for 95% confidence'
>>> low, high = confidence_interval([2, 4, 4, 4, 5, 5, 7, 9], confidence=0.9)
>>> assert math.isclose((low + high) / 2, 5), 'The interval should be centered at the sample mean'
>>> assert math.isclose(high - low, 2 * 1.6448536269514722 * math.sqrt(32 / 7 / 8), rel_tol=1e-6), 'Wrong interval width for 90% confidence (did you use ddof=1?)'
```


Tests for divide_alphabet (points: None)
```
>>> def compare_mappings(m1, m2):
...     assert set(m1.keys()) == set(m2.keys()), 'Mapping keys do not match'
...     for key in m1:
...         assert set(m1[key]) == set(m2[key]), f'Mapping values do not match for key {key}'
...     return True
>>> def test_divide_alphabet():
...     alphabet1 = {'shape', 'shake', 'shame'}
...     guess1 = 'shake'
...     divided1 = divide_alphabet(guess1, alphabet1)
...     ref_divided1 = {(2, 2, 2, 2, 2): {'shake'}, (2, 2, 2, 0, 2): {'shape', 'shame'}}
...     assert compare_mappings(divided1, ref_divided1)
...     alphabet2 = {'three', 'abide', 'drama'}
...     guess2 = 'abide'
...     divided2 = divide_alphabet(guess2, alphabet2)
...     ref_divided2 = {(2, 2, 2, 2, 2): {'abide'}, (0, 0, 0, 0, 2): {'three'}, (1, 0, 0, 1, 0): {'drama'}}
...     assert compare_mappings(divided2, ref_divided2)
>>> test_divide_alphabet()
```


Tests for expected_remaining (points: None)
```
>>> assert math.isclose(expected_remaining([1 / 3, 2 / 3], 3), 5 / 3), 'Wrong answer for the "shake" example'
>>> assert math.isclose(expected_remaining([1.0], 7), 7), 'If every word gives the same pattern, no word is eliminated'
>>> assert math.isclose(expected_remaining([0.25] * 4, 4), 1), 'If every word gives a different pattern, exactly one word remains'
>>> assert math.isclose(expected_remaining([0.1, 0.2, 0.3, 0.4], 10), 3), 'Wrong answer for a non-uniform distribution'
```


Tests for filter_alphabet (points: None)
```
>>> out = filter_alphabet(['shape', 'shake', 'shame'], 'shake', (2, 2, 2, 0, 2))
>>> assert isinstance(out, list), 'filter_alphabet should return a list'
>>> assert sorted(out) == ['shame', 'shape'], 'Wrong words kept for the "shake" example'
>>> assert filter_alphabet(['shape', 'shame'], 'shake', (2, 2, 2, 2, 2)) == [], 'No word should be consistent with this pattern'
>>> def test_filter_alphabet():
...     rng = random.Random(rng_seed)
...     for _ in range(20):
...         secret, guess = rng.choice(possible_words), rng.choice(allowed_guesses)
...         pattern = compute_pattern(guess, secret)
...         out = filter_alphabet(possible_words, guess, pattern)
...         assert secret in out, 'The secret word should always be consistent with its own pattern'
...         expected = {w for w in possible_words if compute_pattern(guess, w) == pattern}
...         assert set(out) == expected, f'Wrong alphabet after guessing {guess} with pattern {pattern}'
>>> test_filter_alphabet()
```


Tests for find_best_guess_optimized (points: None)
```
>>> assert find_best_guess_optimized(possible_words) == best_opener, 'The opening guess should be best_opener'
>>> assert find_best_guess_optimized(['price']) == 'price', 'With 1 word left, guess it'
>>> assert find_best_guess_optimized(['price', 'pride']) in ['price', 'pride'], 'With 2 words left, guess one of them'
>>> small = ['price', 'pride', 'prime']
>>> assert find_best_guess_optimized(small) == find_best_guess(small, small), 'With 3 words left, pick the best guess among them'
>>> large = [w for w in possible_words if compute_pattern('soare', w) == (0, 0, 0, 0, 0)]
>>> assert find_best_guess_optimized(large) == find_best_guess(large, allowed_guesses), 'With more than 3 words left, search all allowed guesses'
```


Tests for play_wordle (points: None)
```
>>> assert play_wordle(create_wordle_game(best_opener)) == 1, 'Guessing the secret word right away should take 1 guess'
>>> def test_play_wordle():
...     rng = random.Random(rng_seed)
...     for secret in rng.sample(possible_words, 20):
...         n = play_wordle(create_wordle_game(secret))
...         assert 1 <= n <= 6, f'Took {n} guesses to find "{secret}"'
>>> test_play_wordle()
>>> assert sum(num_guesses) / len(num_guesses) < 4, 'Wordle solver is not efficient enough. Make sure you ran the 300-game simulation above.'
```


Tests for prob_dist (points: None)
```
>>> def test_prob_dist():
...     groups1 = {(2, 2, 2, 2, 2): {'shake'}, (2, 2, 2, 0, 2): {'shape', 'shame'}}
...     dist1 = prob_dist(groups1)
...     assert len(dist1) == len(groups1), 'Number of probabilities should match number of groups'
...     assert all((g >= 0 for g in dist1)), 'All probabilities should be non-negative'
...     assert math.isclose(sum(dist1), 1), 'Probabilities should sum up to 1'
...     ground_truth1 = [1 / 3, 2 / 3]
...     assert all((math.isclose(a, b) for a, b in zip(sorted(dist1), ground_truth1))), 'Probabilities not correct'
...     groups2 = {(2, 2, 2, 2, 2): {'abide'}, (0, 0, 0, 0, 2): {'three'}, (1, 0, 0, 1, 0): {'drama'}}
...     dist2 = prob_dist(groups2)
...     assert len(dist2) == len(groups2), 'Number of probabilities should match number of groups'
...     assert all((g >= 0 for g in dist2)), 'All probabilities should be non-negative'
...     assert math.isclose(sum(dist2), 1), 'Probabilities should sum up to 1'
...     ground_truth2 = [1 / 3, 1 / 3, 1 / 3]
...     assert all((math.isclose(a, b) for a, b in zip(sorted(dist2), ground_truth2))), 'Probabilities not correct'
>>> test_prob_dist()
```


Tests for solver_agreement (points: None)
```
>>> def test_solver_agreement():
...     for guess in ['soare', 'speed', 'qajaq', 'crane', 'fuzzy']:
...         dist = prob_dist(divide_alphabet(guess, possible_words))
...         assert math.isclose(entropy(dist), guess_entropies(possible_words, [guess])[0]), f'Your divide_alphabet/prob_dist disagree with the solver for "{guess}"'
...         assert math.isclose(expected_remaining(dist, len(possible_words)), guess_expected_remaining(possible_words, [guess])[0]), f'Your expected_remaining disagrees with the solver for "{guess}"'
>>> test_solver_agreement()
```
